# <span style="color:teal;">CIS 211 Project 8:  &nbsp; Bears and Fish</span>

##### Due 11:00 P.M. Thursday Mar 16

##### <span style="color:red">Group Members:</span>

Group work is allowed for this project, but groups are exepcted to do more work (see markdown cells at the end of the notebook for suggestions and ideas).  

If you work in a group only one group member should upload this notebook to Canvas.  Edit this cell to include the name and DuckID for each member of the group:

**Name:**

**Name:**

**Name:**

##  <span style="color:teal;">Overview</span> 

The last project this term is a cellular automaton style biological simulation.  

The sytem is described in Chapter 11 of the Miller and Ranum text.  **Read the description of the simulation in the textbook before you start working on this project.**  Your code will be quite a bit different than the code in the book, but the rules of the simulation are the same, and you need to know how objects interact.

The project has three main parts:
* A class named World that defines a 2D grid where cells are accessed according to their row and column coordinates
* Classes for Bear and Fish, the organisms that can inhabit the world
* Two top level functions, `wbf` and `step_system`; the first will make a world and populate it at random with bears and fish, and the second will run a simulation using that world.

We're written `step_system` for you -- it's in the last code cell in this notebook -- but you need to design, implement, and test the classes and the `wbf` function.

##  <span style="color:teal;">Libraries</span> 

Your program will need to use `numpy` (a numeric processing library) and `random` (the builtin random number generator module).  You can import additional modules if you wish; add the import statements to this code cell.

```python
import numpy as np
import random
import turtle
```

##  <span style="color:teal;">Event Log</span> 

A useful debugging technique is to save descriptions of events in a log.  Any code in the simulator, whether it is the top level function or a method in one of your classes, can call a function named `log`, passing it a description of an event that just occurred.  Examples might be 
```
new Fish in cell (x,y)
```
or 
```
Bear in cell (x,y) eats Fish in cell (x,y)
```

The function below uses a global variable named `logging`.  If you want to take advantage of the `log` function:
* edit the cell to set `logging` to True
* add calls to `log` at various points in your methods

Later, when doing large scale simulations, you can turn off logging simply by setting `logging` to False before you call the top level simulation function.

**Note:** &nbsp; you can change the definition of the `log` function however you want, _e.g._ you can have it write log messages to a file or printadditional information.

```python
logging = False

def log(message):
    if logging:
        print(message)
```

##  <span style="color:teal;">World (20 points)</span>

Describe your World class in the following markdown cell, and write the code for the class in the code cell below the markdown cell.

The minimum requiremens for this class are:
* the constructor should be passed the grid size (number of rows and columns)
* all cells should contain `None` when the world is initialized; later they can contain references to Fish or Bear objects
* include a method named `biota` that returns a list of Fish and Bear objects currently in the grid

You should also define the methods that allow you to access a cell or store a value in a cell using Python's indexing operator.  For example, if `w` is a World object, `w[i,j]` should return the object in row `i`, column `j` (which could be None).  You should also be able to store an item in the grid by using `w[i,j]` on the left side of an assignment statement.

You can have additional instance methods, class variables, or class methods.  Make sure you describe any new additions in the documentation.

**Note:** &nbsp; There are no auto-grader tests for the World class; tests for this class will be done as part of the tests for the Bear and Fish classes.

##### <span style=color:red>Documentation</span> 

**Important** Write your documentation in the following markdown cell.  Do not delete or move this cell.

<!-- grade_id: world_doc, points: 5, grade: true, solution: true -->
My world class initializes max parameters for the world, creates a list of created animals, creates a grid of the world, and sets a bear and fish count to 0. The __repr__ function makes a copy of the grid and prints out the same grid with imagwes of the bears and fish and their ID number. The world class has other methods that: Increment counters, set points in the grid, and add/delete things from the list and grid.

##### <span style=color:red>Code</span> 

**Important** Write the definition of your World class in the following code cell. Do not delete or move this cell.

<!-- grade_id: world, points: 15, grade: true, solution: true -->
```python
class World:
    
    def __init__(self, mx, my):
        
        self._maxX = mx
        self._maxY = my
        self._thingList = np.array([None] * (self._maxX * self._maxY))
        self._grid = self._thingList.reshape(self._maxX, self._maxY)
        self.bearCount = 0
        self.fishCount = 0
            
        self.Bear_pic = "\U0001F43B"
        self.Fish_pic = "\U0001F41F"
        
        self.wturtle = turtle.Turtle()
        self.wscreen = turtle.Screen()
        self.wscreen.setworldcoordinates(0, 0, self._maxX - 1, self._maxY - 1)
        self.wturtle.hideturtle()
        
    def __repr__(self):
        
        printList = np.array([None] * (self._maxX * self._maxY))
        printGrid = self._thingList.reshape(self._maxX, self._maxY)

        
        for i in range(self._maxX):
            for j in range(self._maxY):
                if self._grid[i, j] == None:
                    pass
                elif isinstance(self._grid[i, j], Fish):
                    printGrid[i, j] = self.Fish_pic + str(self._grid[i, j].fish_id)
                elif isinstance(self._grid[i, j], Bear):
                    printGrid[i, j] = self.Bear_pic + str(self._grid[i, j].bear_id)
                    
        return "%s" % (printGrid)
        
    def draw(self):
        
        self.wscreen.tracer(0)
        self.wturtle.forward(self._maxX - 1)
        self.wturtle.left(90)
        self.wturtle.forward(self._maxY - 1)
        self.wturtle.left(90)
        self.wturtle.forward(self._maxX - 1)
        self.wturtle.left(90)
        self.wturtle.forward(self._maxY - 1)
        self.wturtle.left(90)
        
        for i in range(self._maxY - 1):
            self.wturtle.forward(self._maxX - 1)
            self.wturtle.backward(self._maxX - 1)
            self.wturtle.left(90)
            self.wturtle.forward(1)
            self.wturtle.right(90)
        
        self.wturtle.forward(1)
        self.wturtle.right(90)
        
        for i in range(self._maxX - 2):
            self.wturtle.forward(self._maxY - 1)
            self.wturtle.backward(self._maxY - 1)
            self.wturtle.left(90)
            self.wturtle.forward(1)
            self.wturtle.right(90)
            
        self.wscreen.tracer(1)
        
    def biota(self):
        
        animal_list = []
        for animal in self._thingList:
            if animal != None:
                animal_list.append(animal)
        return animal_list
        
    def __getitem__(self, x, y):
        
        return self._grid[x, y]
    
    def __setitem__(self, x, y, value):
        
        self._grid[x, y] = (value)
        
    def freezeWorld(self):
        
        self.wscreen.exitonclick()
        
    def addThing(self, thing, x, y):
        
        thing.setX(x)
        thing.setY(y)
        if isinstance(thing, Fish):
            self.incFish()
            self.__setitem__(x, y, thing)
        else:
            self.incBear()
            self.__setitem__(x, y, thing)
        thing.appear()
        
    def delThing(self, thing):
        
        thing.hide()
        self._grid[thing.getX()][thing.getY()] = None
        if isinstance(thing, Bear):
            thing._world.decBear()
        else:
            thing._world.decFish()
        
    def moveThing(self, oldx, oldy, newx, newy):
        
        self._grid[newy][newx] = self._grid[oldy][oldx]
        self._grid[oldy][oldx] = None
        
    def getMaxX(self):
        
        return self._maxX
    
    def getMaxY(self):
        
        return self._maxY
    
    def liveALittle(self): #chooses a random fish or bear to live during that time period
        
        if self.thingList != []:
            thing = random.randrange(len(self.thingList))
            randomthing = self._thingList[thing]
            randomthing.liveALittle()
            
    def emptyLocation(self, x, y):
        
        if self._grid[x, y] == None:
            return True
        else:
            return False
        
    def lookAtLocation(self, x, y):
        
        return self._grid[x, y]
    
    def getNumBear(self):
        
        return self.bearCount
    
    def getNumFish(self):
        
        return self.fishCount
    
    def incBear(self):
        
        self.bearCount += 1
                    
    def decBear(self):
        
        self.bearCount -= 1
        
    def incFish(self):
        
        self.fishCount += 1
        
    def decFish(self):
        
        self.fishCount -= 1
```

## <span style="color:teal;">Fish and Bears</span>

To run a simulation we need to add a random collection of animals to the world.  The two types of animals in this simulation are fish and bears, and you will write class definitions named Fish and Bear that implement the behaviors of the animals.

### Fish Class

During the simulation a Fish object needs to behave as follows:

(1) Fish are susceptible to overcrowding:  if there are fish in 2 or more neighboring cells the fish dies (it's removed from the simulation)

(2) A fish can reproduce if it has been alive for a certain number of time steps: a random neighboring cell is chosen, and if that cell is empty, a new fish is placed in that cell

(3) A fish can move to another cell:  it picks a random direction, and if the neighboring cell in that direction is unoccupied the fish moves there

The constructor for the Fish class will be passed a reference to a World object and a location, in the form of a tuple with a row number and column number (the object needs to know its location so it can look for other objects in neighboring cells).

The class should include the following methods:
* `live` implements rules 1 and 2 shown above
* `move` implements rule 3
* `location` returns the current grid location (row and column) of the object

Define a class variable named `breed_interval` to specify how many time steps a fish must be alive before it reproduces; the initial value for this variable is 12.


### Bear Class

During the simulation a Bear object needs to behave as follows:

(1) A bear looks for fish in each adjacent cell; if it finds one or more fish it eats one at random 

(2) If a bear has not eaten for certain number of time steps it dies (it's removed from the simulation)

(3) A bear can reproduce if it has been alive for a certain number of time steps: a random neighboring cell is chosen, and if that cell is empty, a new bear is placed in that cell

(4) A bear can move to another cell:  it picks a random direction, and if the neighboring cell in that direction is unoccupied the bear moves there

Define a class variable named `breed_interval` to specify how many time steps a bear must be alive before it reproduces; the initial value for this variable is 8.  Define another class variable named `survive_without_food` to be the number of time steps a bear can live before it dies from starvation; the initial value for this variable is 10.

The constructor will be passed a reference to a World object and a location, in the form of a tuple with a row number and column number (the object needs to know its location so it can look for other objects in neighboring cells).

The class should include the following methods:
* `live` implements rules 1, 2, and 3 shown above
* `move` implements rule 4
* `location` returns the current grid location (row and column) of the object

### Animal (Base Class)?

From the descriptions above it should be apparent that fish and bears have some things in common.

For **full credit** on the coding and documnetation portions of this project you should define a class named Animal and use it as the base class for your Fish and Bear classes.  Some things to think about as you design your classes:
* are there behaviors or operations that are common to both and that can be implemented just once in Animal?
* perhaps a behavior or operation can be defined with a default in the Animal class, and then overridden in the derived class?

One strategy you might consider is to write a complete implementation for one class, either Bear or Fish.  After you have debugged the class and it passes its unit tests you'll have a better idea of what to implement in the other class.  Then you can start moving common behaviors to the Animal class while you are writing the other derived class.

You can still receive **partial credit** if you skip the Animal class and simply write completely separate Fish and Bear classes.  None of the unit tests assume there is a class named Animal.

### Details and Hints

**Fish:** &nbsp; The way the simulation is defined a fish might be eaten before the top level simulation calls the `live` method.  Your `live` method should check to make sure the fish is still alive.  The easiest way to do this is to include an instance variable named `_alive` that is set to True when the fish is initialized and set to False when it dies.

You can have additional instance methods, class variables, or class methods.  Make sure you describe any new additions in the documentation.

**Animal:** &nbsp; Keep this in mind when you design the class hierarchy: every object has an attribute named ``__class__`` (with two underscores before and after the name).  It is a reference to the class an object was defined with.  As an example of how to use it, consider what would happen if we want to define `reproduce` in the Animal class so it is inherited by both Fish and Bear classes.  We need to know the value of `breed_interval` in each object's own class.  We can find this value using the expression
```
self.__class__.breed_interval
```
This will be a reference to `Fish.breed_interval` or `Bear.breed_interval`, depending on whether a `reproduce` was called with Fish object or Bear object.


### <span style="color:teal;">Animal (20 points)</span>

##### <span style=color:red>Documentation</span> 

If you implement the Animal class write your documentation in the following markdown cell.  

**Important** Do not delete or move this cell.

<!-- grade_id: animal_doc, points: 10, grade: true, solution: true -->
My Animal base class initializes points on the grid for each bear and fish created, and sets a breed counter to 0. This class has methods that: set and get coordinates, breed, move, and eat.

##### <span style=color:red>Code</span> 

If you implement the Animal class write the definition in the following code cell

**Important:** &nbsp; Do not delete or move this cell.

<!-- grade_id: animal, points: 10, grade: true, solution: true -->
```python
offsetList = [(-1,1), (0,1), (1,1),
              (-1,0),         (1,0),
             (-1,-1), (0,-1), (1,-1)]
        

class Animal:
    
    def __init__(self):
        
        self.turtle = turtle.Turtle()
        self.turtle.up()
        self.turtle.hideturtle()
        
        self.xpos = 0
        self.ypos = 0
        self._alive = True
        self.breedTick = 0
        
    def setX(self, newX):
        
        self.xpos = newX
        
    def setY(self, newY):
        
        self.ypos = newY
        
    def getX(self):
        
        return self.xpos
    
    def getY(self):
        
        return self.ypos
    
    def location(self):
        
        return (self.getX(), self.getY())
    
    def setWorld(self, aworld):
        
        self._world = aworld
        
    def appear(self):
        
        self.turtle.goto(self.xpos, self.ypos)
        self.turtle.showturtle()
        
    def hide(self):
        
        self.turtle.hideturtle()
        self._alive = False
        
    def move(self):
        
        return self.tryToMove()
        
    def tryToBreed(self):
        
        randomOffsetIndex = random.randrange(len(offsetList))
        randomOffset = offsetList[randomOffsetIndex]
        
        nextx = self.xpos + randomOffset[0]
        nexty = self.ypos + randomOffset[1]
        
        while not (0 <= nextx < self._world.getMaxX() and 0 <= nexty < self._world.getMaxY()):
            ranodmOffsetIndex = random.randrange(len(offsetList))
            randomOffset = offsetList[randomOffsetIndex]
            nextx = self.xpos + randomOffset[0]
            nexty = self.ypos + randomOffset[1]
            
        if self._world.emptyLocation(nextx, nexty):
            loc = (nextx, nexty)
            if isinstance(self._world.lookAtLocation(self.xpos, self.ypos), Fish):
                childThing = Fish(self._world, loc)
                self.breedTick = 0
            else:
                childThing = Bear(self._world, loc)
                
    def tryToEat(self):
        
        adjprey = []
        
        for offset in offsetList:
                newx = self.xpos + offset[0]
                newy = self.ypos + offset[1]
                if 0 <= newx < self._world.getMaxX() and 0 <= newy < self._world.getMaxY():
                    if (not self._world.emptyLocation(newx, newy)) and isinstance(self._world.lookAtLocation(newx, newy), Fish):
                        adjprey.append(self._world.lookAtLocation(newx, newy))                
        
        if len(adjprey) > 0:
            randomprey = adjprey[random.randrange(len(adjprey))]
            preyx = randomprey.getX()
            preyy = randomprey.getY()
            
            self._world.delThing(randomprey)
            self.survive_without_food == 0
        else:
            self.starveTick += 1
            
    def tryToMove(self):
        
        randomOffsetIndex = random.randrange(len(offsetList))
        randomOffset = offsetList[randomOffsetIndex]
        
        nextx = self.xpos + randomOffset[0]
        nexty = self.ypos + randomOffset[1]
        
        while not (0 <= nextx < self._world.getMaxX() and 0 <= nexty < self._world.getMaxY()):
            ranodmOffsetIndex = random.randrange(len(offsetList))
            randomOffset = offsetList[randomOffsetIndex]
            nextx = self.xpos + randomOffset[0]
            nexty = self.ypos + randomOffset[1]
            
        if self._world.emptyLocation(nextx, nexty):
            self._world.moveThing(self.getX(), self.getY(), nextx, nexty)
            self.setX(nextx)
            self.setY(nexty)
            self.turtle.goto(self.xpos, self.ypos)
```

### <span style="color:teal;">Fish (20 points)</span>

##### <span style=color:red>Documentation</span> 

**Important:** &nbsp; Write the documentation for your Fish class in the following markdown cell.  Do not delete or move this cell.

<!-- grade_id: fish_doc, points: 5, grade: true, solution: true -->
My Fish class sets two class variables, overcrowd and a breed interval to the given amounts. it uses the worlds addThing method to add the fish and sets an id number for the fish. It also has a live method that tells the given fish to breed, eat and/or move.

##### <span style=color:red>Code</span> 

**Important:** &nbsp; Write the Python code for your Fish class in the following code cell.  Do not delete or move this cell.

<!-- grade_id: fish, points: 7, grade: true, solution: true -->
```python
class Fish(Animal):
    
    _overcrowd = 2
    breed_interval = 12
    
    def __init__(self, world, loc):
        
        super().__init__()
        self.setWorld(world)
        self._world.addThing(self, loc[0], loc[1])
        self.fish_id = self._world.getNumFish()
        
    def live(self):
        
        adjFish = 0
        
        for offset in offsetList:
            newx = self.xpos + offset[0]
            newy = self.ypos + offset[1]
            if 0 <= newx < self._world.getMaxX() and 0 <= newy < self._world.getMaxY():
                    if (not self._world.emptyLocation(newx, newy)) and isinstance(self._world.lookAtLocation(newx, newy), Fish):
                            adjFish += 1
        
        if adjFish >= self._overcrowd:
            self._world.delThing(self)
        else:
            self.breedTick += 1
            if self.breedTick >= self.breed_interval:
                self.tryToBreed()
                
            self.move()
            
    def getOvercrowd(self):
        
        return self._overcrowd
    
    def getBreedInterval(self):
        
        return self.breed_interval
    
    def setOvercrowd(self, x):
        
        self._overcrowd = x
        
    def setBreedInterval(self, x):
        
        self.breed_interval = x
```

##### <span style="color:red">Autograder Tests:</span>

**Important:** &nbsp;  the code cells in this section will be used by `nbgrader` to run automated tests.  Do not move, delete or alter these cells in any way.

<!-- grade_id: fish_test_1, points: 2, grade: true, solution: false -->
```python
# A new world has no objects
w1 = World(5,5)
assert len(w1.biota()) == 0

# After adding a fish there should be one object
f1 = Fish(w1, (2,2))
assert len(w1.biota()) == 1

# Test the location method
assert f1.location() == (2,2)
print(w1._thingList)
print(w1._grid)
w1
```

Output:
```
[None None None None None None None None None None None None
 <__main__.Fish object at 0x0000024DB7EE3940> None None None None None None
 None None None None None None]
[[None None None None None]
 [None None None None None]
 [None None <__main__.Fish object at 0x0000024DB7EE3940> None None]
 [None None None None None]
 [None None None None None]]
[[None None None None None]
 [None None None None None]
 [None None '🐟1' None None]
 [None None None None None]
 [None None None None None]]
```

<!-- grade_id: fish_test_2, points: 2, grade: true, solution: false -->
```python
# Setting breed_interval to 1 should cause a fish to reproduce when live is called
w2 = World(5,5)
Fish.breed_interval = 1
f2 = Fish(w2, (2,2))
f2.live()
print(w2._grid)
assert len(w2.biota()) == 2
# Reset the interval to original value for remaining tests
Fish.breed_interval = 12
w2
```

Output:
```
[[None None None None None]
 [None <__main__.Fish object at 0x0000024DB948D1D0> None None None]
 [None <__main__.Fish object at 0x0000024DB948D710> None None None]
 [None None None None None]
 [None None None None None]]
[[None None None None None]
 [None '🐟2' None None None]
 [None '🐟1' None None None]
 [None None None None None]
 [None None None None None]]
```

<!-- grade_id: fish_test_3, points: 2, grade: true, solution: false -->
```python
# Make three fish, the one in the middle should die from overcrowding
w3 = World(5,5)
f3 = Fish(w3, (2,2))
Fish(w3, (1,1))
Fish(w3, (3,3))
print(w3._grid)
f3.live()
assert len(w3.biota()) == 2
print(w3._grid)
w3
```

Output:
```
[[None None None None None]
 [None <__main__.Fish object at 0x0000024DB9474710> None None None]
 [None None <__main__.Fish object at 0x0000024DB94747F0> None None]
 [None None None <__main__.Fish object at 0x0000024DB9474978> None]
 [None None None None None]]
<ipython-input-8-bd70e2260290> in live(self)
     23 
     24         if adjFish >= self._overcrowd:
---> 25             self._world.delThing(self)
     26         else:
     27             self.breedTick += 1

<ipython-input-6-d52ab0a23b4e> in delThing(self, thing)
    102         thing.hide()
    103         self._grid[thing.getX()][thing.getY()] = None
--> 104         if isinstance(thing, Bear):
    105             thing._world.decBear()
    106         else:

NameError: name 'Bear' is not defined
```

<!-- grade_id: fish_test_4, points: 2, grade: true, solution: false -->
```python
# When a fish moves it should be within one cell of its original location
w4 = World(5,5)
f4 = Fish(w4, (2,2))
print(w4._grid)
f4.move()
print(w4._grid)
r, c = f4.location()
assert (r,c) != (2,2)
assert abs(r-2) <= 1 and abs(c-2) <= 1
w4
```

###  <span style="color:teal;">Bear (20 points)</span>

##### <span style=color:red>Documentation</span> 

**Important:** &nbsp; Write the documentation for your Bear class in the following markdown cell.  Do not delete or move this cell.

<!-- grade_id: bear_doc, points: 3, grade: true, solution: true -->
My bear class is very similar to my fish class, the only difference being that the beat has a survive_without_food counter that can delete the bear from the world by starving.

##### <span style=color:red>Code</span> 

**Important:** &nbsp; Write the Python code for your Bear class in the following code cell.  Do not delete or move this cell.

<!-- grade_id: bear, points: 7, grade: true, solution: true -->
```python
class Bear(Animal):
    
    survive_without_food= 10
    breed_interval = 8
    
    def __init__(self, world, loc):
        
        super().__init__()
        self.setWorld(world)
        self.setX(loc[0])
        self.setY(loc[1])
        self._world.addThing(self, loc[0], loc[1])
        self.starveTick = 0
        self.bear_id = self._world.getNumBear()
        
    def tryToEat(self):
        
        adjprey = []
        
        for offset in offsetList:
                newx = self.xpos + offset[0]
                newy = self.ypos + offset[1]
                if 0 <= newx < self._world.getMaxX() and 0 <= newy < self._world.getMaxY():
                    if (not self._world.emptyLocation(newx, newy)) and isinstance(self._world.lookAtLocation(newx, newy), Fish):
                        adjprey.append(self._world.lookAtLocation(newx, newy))
                        
        if len(adjprey) > 0:
            randomprey = adjprey[random.randrange(len(adjprey))]
            preyx = randomprey.getX()
            preyy = randomprey.getY()
            
            self._world.delThing(randomprey)
            self.survive_without_food == 0
        else:
            self.starveTick += 1
            
    def live(self):
        
        self.breedTick += 1
        if self.breedTick >= self.breed_interval:
            self.tryToBreed()
            
        self.tryToEat()
        
        if self.starveTick == self.survive_without_food:
            self._world.delThing(self)
        else:
            self.move()
            
    def getStarve(self):
        
        return self.starve
    
    def getBreed(self):
        
        return self.breed
    
    def setStarve(self, x):
        
        self.starve = x
        
    def setBreed(self, x):
        
        self.breed = x
```

##### <span style="color:red">Autograder Tests:</span>

**Important:** &nbsp;  the code cells in this section will be used by `nbgrader` to run automated tests.  Do not move, delete or alter these cells in any way.

<!-- grade_id: bear_test_1, points: 2, grade: true, solution: false -->
```python
# Test the Bear constructor and location method
w1 = World(5,5)
b1 = Bear(w1, (1,1))
assert len(w1.biota()) == 1
assert b1.location() == (1,1)
print(w1._grid)
w1
```

Output:
```
[[None None None None None]
 [None <__main__.Bear object at 0x0000024DB9ADBF28> None None None]
 [None None None None None]
 [None None None None None]
 [None None None None None]]
[[None None None None None]
 [None '🐻1' None None None]
 [None None None None None]
 [None None None None None]
 [None None None None None]]
```

<!-- grade_id: bear_test_2, points: 2, grade: true, solution: false -->
```python
# Repeat the reproduction test for Bears
w2 = World(5,5)
Bear.breed_interval = 1
b2 = Bear(w2, (2,2))
b2.live()
assert len(w2.biota()) == 2
Bear.breed_interval = 8
print(w2._grid)
w2
```

Output:
```
[[None None None None None]
 [None None None None None]
 [None None None <__main__.Bear object at 0x0000024DB9AC30F0> None]
 [None <__main__.Bear object at 0x0000024DB947FEF0> None None None]
 [None None None None None]]
[[None None None None None]
 [None None None None None]
 [None None None '🐻1' None]
 [None '🐻2' None None None]
 [None None None None None]]
```

<!-- grade_id: bear_test_3, points: 2, grade: true, solution: false -->
```python
# Make fish for the bear to eat, count the number of objects after eating
w3 = World(5,5)
b3 = Bear(w3, (2,2))
Fish(w3, (1,1))
Fish(w3, (3,3))
print(w3._grid)
b3.live()
assert len(w2.biota()) == 2
w3
```

Output:
```
[[None None None None None]
 [None <__main__.Fish object at 0x0000024DB947FC88> None None None]
 [None None <__main__.Bear object at 0x0000024DB9AD5A58> None None]
 [None None None <__main__.Fish object at 0x0000024DB9AD5AC8> None]
 [None None None None None]]
[[None None None None None]
 [None '🐻1' None None None]
 [None None None None None]
 [None None None '🐟2' None]
 [None None None None None]]
```

```python
# Make fish for the bear to eat, count the number of objects after eating
w3 = World(5,5)
b3 = Bear(w3, (2,2))
Fish(w3, (1,1))
Fish(w3, (3,3))
b3.live()
print(w3._thingList)
print(w3._grid)
assert len(w3.biota()) == 2
w3
```

Output:
```
[None None None None None None <__main__.Fish object at 0x0000024DB9AD5DA0>
 None None None None None None None None None
 <__main__.Bear object at 0x0000024DB949D208> None None None None None None
 None None]
[[None None None None None]
 [None <__main__.Fish object at 0x0000024DB9AD5DA0> None None None]
 [None None None None None]
 [None <__main__.Bear object at 0x0000024DB949D208> None None None]
 [None None None None None]]
[[None None None None None]
 [None '🐟1' None None None]
 [None None None None None]
 [None '🐻1' None None None]
 [None None None None None]]
```

<!-- grade_id: bear_test_4, points: 2, grade: true, solution: false -->
```python
# Setting the survival limit to 1 should cause a Bear to starve 
w4 = World(5,5)
Bear.survive_without_food = 1
b4 = Bear(w4, (2,2))
print(w4._grid)
b4.live()
print(w4._grid)
assert len(w4.biota()) == 0
w4
```

Output:
```
[[None None None None None]
 [None None None None None]
 [None None <__main__.Bear object at 0x0000024DB9AD5E80> None None]
 [None None None None None]
 [None None None None None]]
[[None None None None None]
 [None None None None None]
 [None None None None None]
 [None None None None None]
 [None None None None None]]
[[None None None None None]
 [None None None None None]
 [None None None None None]
 [None None None None None]
 [None None None None None]]
```

```python
Bear.survive_without_food = 10
```

<!-- grade_id: bear_test_5, points: 2, grade: true, solution: false -->
```python
# Repeat the move test for bears
w5 = World(5,5)
b5 = Bear(w5, (2,2))
print(w5._grid)
b5.move()
print(w5._grid)
r, c = b5.location()
assert (r,c) != (2,2)
assert abs(r-2) <= 1 and abs(c-2) <= 1
w5
```

Output:
```
[[None None None None None]
 [None None None None None]
 [None None <__main__.Bear object at 0x0000024DB9AEE5F8> None None]
 [None None None None None]
 [None None None None None]]
[[None None None None None]
 [None None <__main__.Bear object at 0x0000024DB9AEE5F8> None None]
 [None None None None None]
 [None None None None None]
 [None None None None None]]
[[None None None None None]
 [None None '🐻1' None None]
 [None None None None None]
 [None None None None None]
 [None None None None None]]
```

##  <span style="color:teal;">The `wbf` Function (10 points)</span>

Fill in the body of the `wbf` function so it returns a new World object with the specified number of rows and columns and with the specified number of Bear and Fish objects at random locations.

When we grade your project we will call `wbf` to make a World object and then use the main loop (implemented by `step_system`) to run the simulation.

**There is no documentation requirement for the `wbf` function.**

<!-- grade_id: wbf, points: 5, grade: true, solution: true -->
```python
def wbf(nrows, ncols, nbears, nfish):
    '''
    Return a new World object with the specified number of Bear and Fish objects.
    '''
    myworld = World(nrows, ncols)
    myworld.draw()
    
    def random_location():
        x = random.randrange(myworld.getMaxX())
        y = random.randrange(myworld.getMaxY())
        while not myworld.emptyLocation(x, y):
            x = random.randrange(myworld.getMaxX())
            y = random.randrange(myworld.getMaxY())
        return (x, y)
    
    for bear in range(nbears):
        Bear(myworld, random_location())
        
    for fish in range(nfish):
        Fish(myworld, random_location())
        
    return myworld
```

##### <span style="color:red">Autograder Tests:</span>

**Important:** &nbsp;  the code cells in this section will be used by `nbgrader` to run automated tests.  Do not move, delete or alter these cells in any way.

<!-- grade_id: wbf_test, points: 5, grade: true, solution: false -->
```python
w = wbf(10,10,3,12)

dct = { Bear: 0, Fish: 0 }
for x in w.biota():
    dct[x.__class__] += 1

assert dct[Bear] == 3
assert dct[Fish] == 12
```

##  <span style="color:teal;">The `step_system` Function</span>

We've written this function for you.  It will run a single time step of the simulation.  Pass it a World object containing a grid populated with Bear and Fish objects and it will (a) see which animals survive, then (b) allow all the animals to move to a new location.

```python
def step_system(world):
    for x in world.biota():
        x.live()
    for x in world.biota():
        x.move()
```

##  <span style="color:teal;">Experiments (10 points)</span>

Run some experiments with the top level simulation loop and describe the results in the markdown cell below.  Some ideas of things to try:
* The settings for the Bear class `breed_interval` and `survive_without_food` variables come from the textbook.  Will the world ever run out of bears with these settings?
* Change the settings so `Bear.breed_interval` is larger than `Bear.survive_without_food`.  How does that change the outcome?
* Set the `breed_interval` counter for the Fish class to a smaller number (e.g. 4) so the world has more fish.  What effect does that have?
* Write a function that runs the simulation for a specified number of generations, or until there are no more objects left in the grid.  What combination of parameters leads to the largest number of time steps before the simulation halts?

-It is possible that the world will run out of bears. It just depends on the random location they move to each time. 

-It means that over time, less bears will be reproduced because they must stay alive longer to reproduce.

-The only effect this has is that more fish are breeding and therefore more bears are eating, survivng longer, and reproducing more.

```python
w = wbf(10,10,3,12)
num_generations = 10
for generation in range(num_generations):
    while len(w._thingList) > 0:
        step_system(w)
```

<!-- grade_id: experiments, points: 10, grade: true, solution: true -->
Documentation above

##  <span style="color:teal;">Projects for Groups and/or Extra Credit</span>

Here are some suggestions for ways to extend the simulation.  We will consider other types of extensions -- send a request to `conery@uoregon.edu` with your proposal.

Groups with three people should implement two extensions.  Groups with two people can choose either extension, or do both for extra credit.

* Implement the Plant class described in Section 10.7 of the textbook, and modify the Fish class so fish eat plants and die if they don't find enough food.

* If you implement Plant, how does that affect the class hierarchy?  Is there an even more general class called Organism, with Plants and Animals subtypes of organism?

* Experiment with data structures: make a second version of the World class, but use a list-of-lists approach to making the grid.  Which is more efficient, a list of lists or a numpy array?  Which scales better when larger worlds are used in the simulation?

* Implement a GUI using `tkinter` that is similar to the Solar System GUI.  A canvas should display the world along with images for the Bear and Fish objects (you can download `Bear.gif` and `Fish.gif` from Canvas).  Use a spinbox or text entry box to specify the number of Bear and Fish objects and the number of time steps to run.  Include a Run button to start the simulation.


```python
class Plant():
    
    breed_interval = 5
    
    def __init__(self, world, location):
        
        self.setWorld(world)
        self._world.addThing(self, loc[0], loc[1])
        
    def live(self):
        
        self.breedTick +=1
        if self.breedTick >= self.breed_interval:
            self.tryToBreed()
```

```python

```